# Assignment 1.4 Pandas
Pandas is an open source Python package providing fast, flexible, and expressive data structures and data manipulation designed to make working with “relational” or “labeled” data both easy and intuitive. It is the fundamental high-level building block for doing practical, real world data analysis in Python.
https://pandas.pydata.org/

<!-- grade_id: cell-462d89bbf25b800a, points: 0, grade: false, solution: true -->
```python
## Import numpy and pandas
import numpy as np
import pandas as pd

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-1-96b35f1be191> in <module>()
      4 
      5 # YOUR CODE HERE
----> 6 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-ab515cd97836235b, points: 1, grade: true, solution: false -->
```python

```

The two primary data structures of pandas are 
- Series (1-dimensional) and 
- DataFrame (2-dimensional)

Read this [10 min intro to get started](https://pandas.pydata.org/pandas-docs/stable/10min.html)

### Data Loading

You are given two csv files.  
- 'sales1.csv'
- 'sales2.csv'

Load both CSV as dataframes and Investigate the data and its types

<!-- grade_id: cell-15bfe6e4718647be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sales_1 = pd.read_csv('sales1.csv')
sales_2 = pd.read_csv('sales2.csv')
sales_1.head()
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-2-be93cfd6483d> in <module>()
      3 sales_2 = pd.read_csv('sales2.csv')
      4 sales_1.head()
----> 5 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-3c85ac1eab5f8b19, points: 0, grade: false, solution: true -->
```python
# Investigate the data 
# YOUR CODE HERE

sales_1.describe()
sales_1.dtypes
sales_1.head()
sales_1.tail(5)
sales_1.count()

sales_2.describe()
sales_2.dtypes
sales_2.head()
sales_2.tail(5)
sales_2.count()
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-3-83a950248260> in <module>()
     13 sales_2.tail(5)
     14 sales_2.count()
---> 15 raise NotImplementedError()

NotImplementedError:
```

### Merge Dataframes
Merge the two dataframes into a single dataframe "sales" based on order id such that new data consist of common order ID and all the other columns.
Also analyse the data in different columns.
Use [Merge function](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.merge.html)

<!-- grade_id: cell-82ee3339bf5f4570, points: 0, grade: false, solution: true -->
```python
sales = pd.merge( sales_1, sales_2, on="Order ID", how="outer")
# YOUR CODE HERE
sales.count()
sales.describe()
sales.dtypes
sales.head(5)
sales.tail(5)
sales[sales["Order ID"]==686800706]["Country"].values
sales[sales["Order ID"]==686800706]["Item Type"].values
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-4-946c55105c0b> in <module>()
      8 sales[sales["Order ID"]==686800706]["Country"].values
      9 sales[sales["Order ID"]==686800706]["Item Type"].values
---> 10 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-52d67c318540c515, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==686800706]["Country"].values == "Libya")
assert(sales[sales["Order ID"]==686800706]["Item Type"].values == "Cosmetics")
```

### Data Cleaning
If you look at the data we can see some issues,
- Some of the Item Type start with " * " we want to remove the asterik in the column
- Some columns data are missing indicated by NaN (u can try isnull(), isna() function)
  - If Region or Country data is missing replace it with unknown
  - if other data are missing, remove it from "sales" dataframe and add it in a "missing" dataframe 

Reference for this section:
- [Working with missing data](https://pandas.pydata.org/pandas-docs/stable/missing_data.html)
- [Python Lambda Expressions](https://docs.python.org/3/tutorial/controlflow.html#lambda-expressions)
- [Pandas Apply Function](https://pandas.pydata.org/pandas-docs/version/0.18/generated/pandas.Series.apply.html)
- [Pandas FillNa](https://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.fillna.html)

<!-- grade_id: cell-36c46b4d53bf011c, points: 0, grade: false, solution: true -->
```python
# Investigate dataframe for Null (NA/NAN) and count the number of rows with NA/NAN
null_values = sales.isnull()
NA_count = sales.isna().sum(axis = 0)
NA_count
null_count  = sales[sales.isnull().any(axis = 1)].shape[0]
null_count
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-6-e5125f324337> in <module>()
      5 null_count  = sales[sales.isnull().any(axis = 1)].shape[0]
      6 null_count
----> 7 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-c3f7ee60f7aec5bf, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-04b3a523548b2603, points: 0, grade: false, solution: true -->
```python
# remove "*", NaN, add missing sales data to "missing" dataframe 
# sales = sales['Country'].fillna('Unknown')    ## Do not use this method as it overwrites the dataframe


sales['Item Type'] = sales['Item Type'].apply(lambda x:str(x).replace('*',''))
sales = sales.fillna({'Country': 'unknown',  'Region' : 'unknown'})

missing = sales[sales.isna().any(axis=1)]
# missing = pd.DataFrame(x) ##creating new dataframe with null values
# missing

# sales = sales.dropna(axis = 0, how ='any')  ## Removing null values from existing dataframe
# sales = sales.reset_index(drop=True)

# raise NotImplementedError()
```

```python
sales = sales[sales.notnull().all(axis = 1)]
sales

```

Output:
```
Item Type  Order Date   Order ID   Ship Date  Units Sold  \
0          Cosmetics  10/18/2014  686800706  10/31/2014      8446.0   
2          Baby Food  10/31/2016  246222341   12/9/2016      1517.0   
3             Cereal   4/10/2010  161442649   5/12/2010      3322.0   
4             Fruits   8/16/2011  645713555   8/31/2011      9845.0   
5             Cereal  11/24/2014  683458888  12/28/2014      9528.0   
6             Cereal    3/4/2015  679414975   4/17/2015      2844.0   
7            Clothes   5/17/2012  208630645   6/28/2012      7299.0   
8         Vegetables   1/29/2015  266467225    3/7/2015      2428.0   
9         Vegetables  12/24/2013  118598544   1/19/2014      4800.0   
10           Clothes  12/29/2015  451010930   1/19/2016      3012.0   
11            Snacks   2/27/2010  220003211   3/18/2010      2694.0   
12         Household  11/17/2016  702186715  12/22/2016      1508.0   
31   Office Supplies   6/30/2015  499009597    7/9/2015      6884.0   
32            Fru
```

<!-- grade_id: cell-7bc1b916b4472109, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==267066323]["Item Type"].values == "Cereal")
assert(sales[sales["Order ID"]==267066323]["Region"].values == "unknown")
```

Find the number of unique items available.

<!-- grade_id: cell-588bcfca26c28277, points: 0, grade: false, solution: true -->
```python
# NumItems = sales['Item Type'].nunique()
NumItems = len(sales['Item Type'].unique())
NumItems
# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-10-9f730129bc9f> in <module>()
      3 NumItems
      4 # YOUR CODE HERE
----> 5 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-51b5cda4685cd3d0, points: 1, grade: true, solution: false -->
```python

```

### Grouping and Analysis 
Find out the statistical summary for ‘Unit Cost’ column for each Items.
Then extract the count, mean and std for "Cereal","Fruits","Vegetables","Baby Food" and "Meat" as a dataframe "Kitchen"
References:
- [Pandas Indexing](https://pandas.pydata.org/pandas-docs/stable/indexing.html)
- [Groupby](https://pandas.pydata.org/pandas-docs/stable/groupby.html)
- [Loc](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.DataFrame.loc.html)
- [filter](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.filter.html)

```python
# x = sales.groupby('Item Type').describe().transpose()
# x = sales.groupby('Item Type').describe() ## Gives the statistical overview of all when data is grouped by food items



```

```python
# x = sales['Item Type'].unique()
# for items in x:
#     print(items)     ## To display distince items

```

```python
##### sales['Unit Cost'].describe
# sales.nunique('Item Type')
# by_cost = sales.groupby("Unit Cost")
# by_cost.describe().transpose()
# sales.unique('Item Type')
# x = sales.filter(items=['Item Type'])
# sales['Item Type'].unique().tolist()
sales.groupby(['Item Type'], as_index = False)["Unit Cost"].mean()

```

Output:
```
Item Type  Unit Cost
0         Baby Food     159.42
1         Beverages      31.79
2            Cereal     117.11
3           Clothes      35.84
4         Cosmetics     263.33
5            Fruits       6.92
6         Household     502.54
7              Meat     364.69
8   Office Supplies     524.96
9     Personal Care      56.67
10           Snacks      97.44
11       Vegetables      90.93
```

<!-- grade_id: cell-d5b95b853fd96b60, points: 0, grade: false, solution: true -->
```python
food_items = sales.groupby('Item Type')
food_summary = food_items['Unit Cost'].describe()   ## Gives statistical summary for "unit cost" for each items i.e when we groupby items. Groupby groups the data according to
##value we define. Here, we have grouped out data by Item types, hence the data is shown with respect to item types

Kitchen = food_summary.filter(items = ['count', 'mean', 'std'])
Kitchen = Kitchen.loc[['Cereal', 'Fruits', 'Vegetables', 'Baby Food', 'Meat']]
Kitchen

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
count    mean           std
Item Type                              
Cereal       77.0  117.11  4.291212e-14
Fruits       68.0    6.92  1.789564e-15
Vegetables   95.0   90.93  2.857249e-14
Baby Food    85.0  159.42  2.859039e-14
Meat         77.0  364.69  5.721617e-14
```

<!-- grade_id: cell-986e659a273a364c, points: 1, grade: true, solution: false -->
```python
assert(Kitchen.loc["Cereal"]["count"]==77)
```

<!-- grade_id: cell-346f31cb4ea47c75, points: 0, grade: false, solution: true -->
```python
#Generate the new column ‘Total Profit’ for each orders (Total Profit = Total Revenue - Total Cost)
# YOUR CODE HERE
sales['Total Profit'] = sales['Total Revenue'] - sales['Total Cost']

# raise NotImplementedError()

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-14-81f0d4866903> in <module>()
      3 sales['Total Profit'] = sales['Total Revenue'] - sales['Total Cost']
      4 
----> 5 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-6419595bc71f6a09, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==246222341]["Total Profit"].values == 145419.62)
```

<!-- grade_id: cell-aae9d532c4dabf17, points: 0, grade: false, solution: true -->
```python
# How many data sets are there with profit greater than 60%

# sales['Total Profit'] = sales.apply(lambda row: ((row['Total Revenue'] - row['Total Cost'])/ row['Total Cost'])*100, axis=1)  ##Calculating profit percent
# sales.apply(lambda x: (x['Total Profit']/x['Total Cost'])*100)

sales['Profit percent'] = sales['Total Profit']/sales['Total Cost']*100
ans = (sales['Profit percent']>=60).count()
ans
# raise NotImplementedError()

```

Output:
```
404
```

<!-- grade_id: cell-c053463e75781632, points: 1, grade: true, solution: false -->
```python

```